# Learning Activity 9

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from numpy import sin, cos, tan, pi
import robotics
from robotics import Rz, Ry, Rz, sind, cosd, DH
rtd = 180/pi
dtr = pi/180
```

```python
def myarraytest(in1, in2):
    mydiff = np.array(in1) - np.array(in2)
    mymax = np.max(np.abs(mydiff))
    return mymax < 1e-6
```

## Problem 1

<img src="https://faculty.gvsu.edu/kraussry/445_images/learning_activities/inverse_kin_sketches/inv_kin_simple_robot_1.png" width=400px>

### DOF and Problem Statement

- how many degrees of freedom does this robot have?
- what is required for a valid problem statement?

```python
A = 4
B = 3
C = 2
```

Solve the inverse kinematics to place the tip at

$$^0P_{tip} = \begin{bmatrix} 2.328 \\ 0.7252 \\ 5.92 \\ 1 \end{bmatrix}$$

```python
P_tip_0 = np.array([2.32834326, 0.72519817, 5.91961058, 1])
```

**Key Question:** 

- is this point in the reachable workspace?

**Note:** be sure to use the variables `th1_p1` and `th2_p1` for your answers to problem 1.

### Recommended Procedure

- draw a top view in frame 0
    - use the top view to find $\theta_1$
- calculate $^1P_{tip}$ by using $\theta_1$ to find $^0_1\mathbf{T}$
- draw a side view in frame 1
- find $\theta_2$ based on the side view

## Test

<!-- grade_id: cell-ab7f9375f42d3eab, points: 20, grade: true, solution: false -->
```python
P_tip_1 = np.array([B+C*cosd(th2_p1),0,C*sind(th2_p1),1])
T01 = robotics.DH(0,0,th1_p1,A)
P_tip_0_test = np.dot(T01,P_tip_1)
assert myarraytest(P_tip_0,P_tip_0_test)
P_tip_0_test
```

Output:
```
array([2.32834326, 0.72519817, 5.91961058, 1.        ])
```

## Problem 2

<img src="https://faculty.gvsu.edu/kraussry/445_images/learning_activities/inverse_kin_sketches/inv_kin_2.jpeg" width=400px>

```python
A2 = 8
B2 = 2
C2 = 5
D2 = 4
```

### Problem Statement

Solve the inverse kinematics to place the tip at

$$^0P_{tip} = \begin{bmatrix} 8.269 \\ -2.371 \\ 13.42\end{bmatrix}$$

```python
P_tip_0_p2_original = np.array([ 8.26927707, -2.37117704, 13.41775726,  1.        ])
```

**Note:** use the variables `th1_p2`, `th1_p2`, and `th1_p2` for your answer to problem 2.

### Recommended Procedure

- draw a top view in frame 0
    - use the top view to find $\theta_1$
- calculate $^1P_{tip}$ by using $\theta_1$ to find $^0_1\mathbf{T}$
- draw a side view in frame 1
- find $\theta_2$ and $\theta_3$ based on two-link inverse kinematics based on the side view

## Test

<!-- grade_id: cell-52d485c77e5d46f4, points: 30, grade: true, solution: false -->
```python
x1 = B2 + C2*cosd(th2_p2) + D2*cosd(th2_p2+th3_p2)
z1 = C2*sind(th2_p2) + D2*sind(th2_p2+th3_p2)
P_tip_1_p2 = np.array([x1,0,z1,1])
T01_p2 = robotics.DH(0,0,th1_p2,A2)
P_tip_0_p2 = np.dot(T01_p2,P_tip_1_p2)
P_tip_0_p2
```

```python

```