# Learning Activity 7: Intro to DH

```python
%matplotlib inline
import matplotlib.pyplot as plt
import numpy as np
from numpy import sin, cos, tan, pi
import robotics
from robotics import Rx, Ry, Rz, sind, cosd, DH
rtd = 180/pi
dtr = pi/180
```

```python
def myarraytest(in1, in2):
    mydiff = np.array(in1) - np.array(in2)
    mymax = np.max(np.abs(mydiff))
    return mymax < 1e-6
```

## Problem 1

- Attach the frames necessary to model the robot below using the DH convention
- fill out a table of the DH parameters
- find $^0P_{tip}$ given the values:
    - $L_1$ = 9 cm
    - $L_2$ = 8 cm
    - $L_3$ = 7 cm
    - $\theta_1$ = 20$^\circ$
    - $\theta_2$ = -40$^\circ$
    - $\theta_3$ = 15$^\circ$
 

<img src="https://faculty.gvsu.edu/kraussry/445_images/learning_activities/dh/dh1_3_link_planar.png" width=400px>

```python
L1 = 9
L2 = 8
L3 = 7
th1 = 20
th2 = -40
th3 = 15
```

**Note:** Use the variable `P_tip_0_1` for your final answer.

### Test

<!-- grade_id: cell-4405c836394f024a, points: 10, grade: true, solution: false -->
```python
x_check = L1*cosd(th1)+L2*cosd(th1+th2)+L3*cosd(th1+th2+th3)
y_check = L1*sind(th1)+L2*sind(th1+th2)+L3*sind(th1+th2+th3)
P_tip_test_1 = np.array([x_check,y_check,0,1])
assert myarraytest(P_tip_0_1,P_tip_test_1)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-be4995314626> in <module>
----> 1 x_check = L1*cosd(th1)+L2*cosd(th1+th2)+L3*cosd(th1+th2+th3)
      2 y_check = L1*sind(th1)+L2*sind(th1+th2)+L3*sind(th1+th2+th3)
      3 P_tip_test_1 = np.array([x_check,y_check,0,1])
      4 assert myarraytest(P_tip_0_1,P_tip_test_1)

NameError: name 'L1' is not defined
```

## Problem 2

- Attach the frames necessary to model the robot below using the DH convention
- fill out a table of the DH parameters
- find $^0P_{tip}$ given the values:
    - $A$ = 10 cm
    - $B$ = 8 cm
    - $C$ = 6 cm
    - $\theta_1$ = 30$^\circ$
    - $\theta_2$ = 20$^\circ$
    - $\theta_3$ = -15$^\circ$

<img src="https://faculty.gvsu.edu/kraussry/445_images/learning_activities/dh/dh2.jpeg" width=500px>

```python
A = 10
B = 8
C = 6
th1 = 30
th2 = 20
th3 = -15
```

### Test

<!-- grade_id: cell-d646f5a097131b2c, points: 10, grade: true, solution: false -->
```python
assert myarraytest(P_tip_0_3,[21.33718613, 10.63401945, -0.53112796,  1.        ])
print("nice work")
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-1-39e635918d76> in <module>
----> 1 assert myarraytest(P_tip_0_3,[21.33718613, 10.63401945, -0.53112796,  1.        ])
      2 print("nice work")

NameError: name 'myarraytest' is not defined
```

```python

```

```python

```

## Problem 3

- draw the schematic and fill in the DH table

<img src="https://faculty.gvsu.edu/kraussry/445_images/learning_activities/dh/dh3.jpeg" width=500px>


```python

```